# <span style="color:teal;">CIS 211 Project 7:  SQLite</span>

##### Due 11:00 P.M. Friday Feb 26

This project will give you some experience writing database queries in SQL and connecting to a SQLite database from Python.

### <span style="color:teal">Setup</span>

Download the Sakila database (the file named `sakila211.db`) and move it to the same folder as this IPython notebook.

### <span style="color:teal">Executing Queries in an IPython Notebook</span>

To solve the problems in the first two sections of this project you need to write a SQL query based on the tables in the Sakila211 database.  

To execute a query from an IPython Notebook write the query in quotes as part of a shell command:
```
! sqlite3 sakila211.db "..."
```

For example, to answer the question "What are the names and ratings of the films" your code cell would have
```
! sqlite3 sakila211.db "select title, rating from film"
```

Some of the queries will produce a lot of output.  You can add a `LIMIT` clause to any query if you want to reduce the size of the output.

**Note:** &nbsp; Your grade will not depend on the output format.  If want to have a nicer looking output that includes column headers create a file named `.sqliterc` in your home folder and add these two lines:
```
.header on
.mode column
```

```python
! sqlite3 sakila211.db ".tables"
```

## <span style="color:teal">Part 1:  &nbsp; Queries</span>

The following questions can all be answered using records from a single table.

(1) What are the first and last names of all the customers?

<!-- grade_id: query_1, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select first_name,last_name from customer"
```

(2) What are the e-mail addresses of customers whose first name is “Willard”?

<!-- grade_id: query_2, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select email from customer where first_name = 'Willard'"
```

(3) How many customers are there?

<!-- grade_id: query_3, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select count(last_name) from customer"
```

(4) How many customers shop at store number 1?

<!-- grade_id: query_4, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select count(last_name) from  customer where store_id='1'"
```

(5) How much does it cost to rent the film named “Virtual Spoilers”?

<!-- grade_id: query_5, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select rental_rate from  film where title='Virtual Spoilers'"
```

(6) Do any films have the word “Princess” in the title?

<!-- grade_id: query_6, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select count(title) from  film where title like'%Princess%'"
```

(7) What are the titles of the films that are longer than 180 minutes?

<!-- grade_id: query_7, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select title from  film where length>'180'"
```

(8) How many films have a “G” rating and are less than 60 minutes long?

<!-- grade_id: query_8, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select count(title) from  film where rating='G' and length>'60'"
```

(9) What is the maximum replacement cost for any film?

<!-- grade_id: query_9, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select max(replacement_cost) from  film "
```

(10) Print a table that lists the different types of ratings and the number of films that have that rating.

<!-- grade_id: query_10, points: 2, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select rating, count(title) from  film group by rating"
```

## <span style="color:teal">Part 2: &nbsp; Joins</span>

The following queries all require a join of two or more tables. As a hint for how to create the query the table names are shown next to each a question.

(11) What are the film IDs of movies starring Jude Cruise?  &nbsp; &nbsp; [`actor`, `film_actor`]

<!-- grade_id: query_11, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select film_id as film_id_acted_by_Jude from actor join film_actor using(actor_id) where first_name='Jude'"
```

(12) What language was “Ace Goldfinger” filmed in?  &nbsp; &nbsp; [`film`, `language`]

<!-- grade_id: query_12, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select  name as language from film join language  using(language_id) where title='Ace Goldfinger' "
```

(13) How many actors starred in the film named “Scalawag Duck”? &nbsp; &nbsp; [`film`, `film_actor`]

<!-- grade_id: query_13, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select  count(actor_id) as number_of_actors  from film_actor join film using(film_id) where title='Scalawag Duck' "
```

(14) Produce a table that shows how many films each customer has rented.  &nbsp; &nbsp; [`customer`, `rental`]

<!-- grade_id: query_14, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select  customer_id , first_name,last_name, count(customer_id) as number_of_film_they_rent from customer join rental using(customer_id) group by customer_id"
```

Output:
```
-- Loading resources from /Users/Andyfish/.sqliterc

customer_id  first_name  last_name   number_of_film_they_rent
-----------  ----------  ----------  ------------------------
1            Mary        Smith       32                      
2            Patricia    Johnson     27                      
3            Linda       Williams    26                      
4            Barbara     Jones       22                      
5            Elizabeth   Brown       38                      
6            Jennifer    Davis       28                      
7            Maria       Miller      33                      
8            Susan       Wilson      24                      
9            Margaret    Moore       23                      
10           Dorothy     Taylor      25                      
11           Lisa        Anderson    24                      
12           Nancy       Thomas      28                      
13           Karen       Jackson     27                      
14           Bett
```

(15) How many films did the customer named Smith rent?  &nbsp; &nbsp; [`customer`, `rental`]

<!-- grade_id: query_15, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select   count(customer_id) as number_of_Smith_rent  from customer  join rental  using(customer_id) where last_name = 'Smith' group by customer_id"
```

(16) Which customers have not returned films (*i.e.* the return_date field in the rental table is null)?  &nbsp; &nbsp; [`customer`, `rental`]

<!-- grade_id: query_16, points: 3, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select   first_name,last_name, count(customer_id) as number_of_films_did_not_return from customer  join rental  using(customer_id) where return_date is null group by customer_id"
```

(17) Which actors starred in “Splash Gump”?  &nbsp; &nbsp; [`film`, `film_actor`, `actor`]

<!-- grade_id: query_17, points: 4, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select first_name,last_name from film join film_actor using(film_id)join actor  using(actor_id) where title='Splash Gump' group by actor_id  "
```

(18) Which films has Penelope Guiness starred in?  &nbsp; &nbsp; [`film`, `film_actor`, `actor`]

<!-- grade_id: query_18, points: 4, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select title from film join film_actor using(film_id)join actor \
using(actor_id) where last_name='Guiness' group by actor_id  "
```

Output:
```
-- Loading resources from /Users/Andyfish/.sqliterc

title             
------------------
Wizard Coldblooded
United Pilot      
Wind Phantom
```

(19) What are the names of the films rented by the customer named Knight?  &nbsp; &nbsp; [`customer`, `rental`, `inventory`, `film`]

<!-- grade_id: query_19, points: 4, grade: true, solution: true -->
```python
! sqlite3 sakila211.db "select title from film join inventory using(film_id) join rental using(inventory_id) join customer using(customer_id) where last_name='Knight'  "
```

Output:
```
-- Loading resources from /Users/Andyfish/.sqliterc

title        
-------------
Jericho Mulan
Frost Head   
Clerks Angels
Wife Turn    
Expecations N
Lawless Visio
Soup Wisdom  
Heartbreakers
Bikini Borrow
Attacks Hate 
Cupboard Sinn
Princess Gian
Cyclone Famil
Platoon Insti
Knock Warlock
Excitement Ev
Right Cranes 
Hyde Doctor  
Fatal Haunted
Chasing Fight
Falcon Volume
Tomorrow Hust
Half Outfield
Horror Reign 
Borrowers Bed
```

## <span style="color:teal">Part 3: &nbsp; Monthly Report</span>

For this part of the project you will write a report generator that will print a summary of monthly activity for a customer.  The report should show a list of all the films rented during the month along with a total amount due.
This example shows the activity for the customer named Leo Ebert for the month of June, 2005.

<pre>
--- Sakila DVD Rentals ---

Monthly Report for Leo Ebert 

Lonely Elephant                 06/16/2005    $2.99
                   ** late fee  06/24/2005    $2.99
Enemy Odds                      06/18/2005    $4.99
                   ** late fee  06/27/2005    $4.99
Song Hedwig                     06/20/2005    $0.99
                   ** late fee  06/26/2005    $0.99
Wash Heavenly                   06/21/2005    $4.99
Hurricane Affair                06/21/2005    $2.99

Monthly Total:  $25.92
</pre>

To compute the total, the report generator needs to add up all the rental fees, and for any film returned late, add an additional rental fee as the late charge. In the example, Leo rented *Lonely Elephant* on June 16 but didn’t return it until June 24. It was a 3-day rental, so he was charged an additional \$2.99 (it was 5 days late, so we could have charged him \$4.98 in late fees, but we’ll keep it simple and just charge the basic fee as the late fee).

We've broken the problem into separate parts so they can be implemented and tested separately.  Your job is to fill in the code cells so each part passes the auto-grader tests.

**Note:** &nbsp; There is no documentation requirement for this project.

### <span style="color:teal">Database Connection</span>

To simplify the code for this program you can define the database connection as a global variable.  The following code cell imports the libraries you need and defines the database connection.

```python
from sqlite3 import connect

db = connect('sakila211.db')
```

### <span style="color:teal">Dates</span>

Dates in a SQLite database are stored as strings.  The function below uses a Python library named `datetime` to parse a date string and return a string that can be printed as part of the report.  Read the online documentation for the `datetime` module if you want to learn more about date formats and functions that parse and format dates.

**Note:** &nbsp; The format string is a global variable -- you can use this same format in your function that checks to see if a CD is overdue.

```python
from datetime import datetime

date_format = "%Y-%m-%d %H:%M:%S.%f"

def mmddyyyy(date):
    d = datetime.strptime(date, date_format)
    return '{:02d}/{:02d}/{:02d}'.format(d.month, d.day, d.year)
```

```python
assert mmddyyyy('2005-02-26 20:59:35.000') == '02/26/2005'
```

```python
d=mmddyyyy('2005-02-26 20:59:35.000') 

```

### <span style="color:teal">Customer Name</span>

Fill in the definition of the `customer_name` function.  It will be passed a string with a last name, and it should return a list of tuples of the form `(first, last)` with the first and last names of all customers with the same last name.  If there are no customers with that name the list should be empty.

<!-- grade_id: customer_name, points: 4, grade: true, solution: true -->
```python
def customer_names(s):
    
    query = "select first_name, last_name from customer where last_name like ?"
    result = db.execute(query,(s,))
    return result.fetchall()

    
```

```python


```

##### <span style="color:red">Autograder Tests:</span>

**Do not edit, copy, or delete these cells.**

<!-- grade_id: cname_1, points: 2, grade: true, solution: false -->
```python
assert customer_names('ebert') == [('Leo', 'Ebert')]
```

<!-- grade_id: cname_2, points: 2, grade: true, solution: false -->
```python
assert customer_names('smith') == [('Mary', 'Smith')]
```

<!-- grade_id: cname_3, points: 2, grade: true, solution: false -->
```python
assert customer_names('turing') == []
```

### <span style="color:teal">Film List</span>

Fill in the body of the `film_list` function so it returns a list of all CDs rented by a customer during a specified month. 
The parameters that will be passed to the function are the customer's last name (a string) and the month and year (both integers).  Return a list of tuples that have the following values, in the order shown:
* film title
* rental rate (how much it costs to rent the CD)
* rental date (the date the CD was rented)
* return date (the date the CD was returned)
* rental duration (the number of days the CD can be out before a late fee is charged) 

All of the information you need can be retrieved with a single query that joins the rental, inventory, film, and customer tables. Note that this is the same query used to answer Question 19 above, except with different column names.

**Hint:** &nbsp; The dates in the database are strings.  Think about pattern matching when you write the WHERE clause that specifies the date range.


<!-- grade_id: film_list, points: 6, grade: true, solution: true -->
```python
def film_list(name, month, year):
    name=name.title()
    a=[]
    L='{:02d}-{:02d}-{:02d} 00:00:00:000'.format(year, month,1)
    R='{:02d}-{:02d}-{:02d} 23:59:59:999'.format(year,month,31)
    query = "select title, rental_rate, rental_date, return_date, rental_duration \
    from film join inventory using(film_id) join rental using(inventory_id) join customer using(customer_id) where last_name = ? and rental_date > ? and rental_date< ?"
    result= db.execute(query,(name,L,R))
    
    return [i for i in result]
        
```

```python
film_list('Ebert', 6, 2005)

```

Output:
```
[('Lonely Elephant',
  2.99,
  '2005-06-16 20:59:35.000',
  '2005-06-24 15:31:35.000',
  3),
 ('Enemy Odds', 4.99, '2005-06-18 19:04:41.000', '2005-06-27 17:10:41.000', 5),
 ('Song Hedwig',
  0.99,
  '2005-06-20 11:17:20.000',
  '2005-06-26 13:15:20.000',
  3),
 ('Wash Heavenly',
  4.99,
  '2005-06-21 02:47:56.000',
  '2005-06-27 03:16:56.000',
  7),
 ('Hurricane Affair',
  2.99,
  '2005-06-21 10:56:59.000',
  '2005-06-25 16:41:59.000',
  6)]
```

##### <span style="color:red">Autograder Tests:</span>

**Do not edit, copy, or delete these cells.**

<!-- grade_id: film_1, points: 2, grade: true, solution: false -->
```python
leo = film_list('Ebert', 6, 2005)
assert len(leo) == 5
assert leo[0] == (('Lonely Elephant', 2.99, '2005-06-16 20:59:35.000', '2005-06-24 15:31:35.000', 3))
```

<!-- grade_id: film_2, points: 2, grade: true, solution: false -->
```python
mary = film_list('Smith', 8, 2005)
assert len(mary) == 11
assert mary[0] == ('Adaptation Holes', 2.99, '2005-08-01 08:51:04.000', '2005-08-10 12:12:04.000', 7)
```

### <span style="color:teal">Overdue CDs</span>

To decide whether a CD was returned late you need to subtract the rental date from the return date and see if the difference is greater than the rental duration. The easiest way to do this is to use Python’s `datetime` library.

Using a call to `strptime` (see the body of the `mmddyyyy` function defined above) create `datetime` objects for the rental date and the return date. Then subtract the rental date from the return date to get a `timedelta` object:
```
    diff = returned - rented
```
Now you can get attributes of the object named `diff` to tell you how far apart the dates were, for example `diff.days` will tell you the time difference in days.

Fill in the body of the `overdue` function.  The parameters are a rental date, a return date, and the rental duration.  Return TRUE if the CD was not returned or if the number of days it was checked out is greater than the rental duration.

<!-- grade_id: overdue, points: 4, grade: true, solution: true -->
```python
def overdue(rent_date, return_date, duration):
    date_format = "%Y-%m-%d %H:%M:%S.%f"

    
    ren = datetime.strptime(rent_date,date_format)
    ret = datetime.strptime(return_date,date_format)
    diff=ret-ren
    
    if diff.days >duration :
        return True
    
    
```

```python
overdue('2005-06-16 20:59:35.000', '2005-06-24 15:31:35.000', 3)
```

Output:
```
True
```

<!-- grade_id: overdue_1, points: 2, grade: true, solution: false -->
```python
assert overdue('2005-06-16 20:59:35.000', '2005-06-19 20:59:35.000', 2)
```

<!-- grade_id: overdue_2, points: 2, grade: true, solution: false -->
```python
assert not overdue('2005-06-16 20:59:35.000', '2005-06-17 20:59:35.000', 2)
```

<!-- grade_id: overdue_3, points: 2, grade: true, solution: false -->
```python
assert not overdue('2005-06-16 20:59:35.000', '2005-06-19 20:59:34.000', 2)
```

### <span style="color:teal">Report Generator</span>

Fill in the body of the function named `sakila_report` so it prints a monthly report.  The arguments passed to the function will be a customer's last name (a string) and a month and year (both integers).

If the customer is not in the database, or if there are two customers with the same last name, return None.  Otherwise print a list of films and late fees and the total amount due.

You can format the report however you would like, as long as the following items appear somewhere in the output:
* every fee should be printed with a dollar sign and an amount, for example `$2.99` or `25.92`
* the total amound due should be printed on a line that includes the word `"Total"` followed by a dollar amount on the same line

<!-- grade_id: report, points: 10, grade: true, solution: true -->
```python
def sakila_report(customer, month, year):
    a= customer_names(customer)
    if len(a) == 0 or len(a)>1:
        return None
    rental = film_list(customer,month,year)
    total_payment =0
    total = 0
    print("---Sakila DVD Rentals---")
    print()
    dis = 'monthly report for {} {}'.format(a[0][0],a[0][1])
    for i in rental:
        total += i[1]
        rental_date = mmddyyyy(i[2])
        dis +='\n \n${} <{}> ranted at :{}'.format(i[1],i[0],rental_date)
        
        if overdue(i[2],i[3],i[4])== True:
            total_payment += i[1]
            dis += '\nOver due \n *add late fee ${}'.format(i[1])
    total_payment = total_payment +total
    dis +='\n Monthly Total: ${0:.2f}'.format(total_payment)
    print(dis)
```

##### <span style="color:red">Tests</span>

You can use code cells in this section (and add more if you wish) to test your functions.

```python

sakila_report('Ebert', 6, 2005)
```

Output:
```
---Sakila DVD Rentals---

monthly report for Leo Ebert
 
$2.99 <Lonely Elephant> ranted at :06/16/2005
Over due 
 *add late fee $2.99
 
$4.99 <Enemy Odds> ranted at :06/18/2005
Over due 
 *add late fee $4.99
 
$0.99 <Song Hedwig> ranted at :06/20/2005
Over due 
 *add late fee $0.99
 
$4.99 <Wash Heavenly> ranted at :06/21/2005
 
$2.99 <Hurricane Affair> ranted at :06/21/2005
 Monthly Total: $25.92
```

```python
check_report(res.stdout, 8, '$25.92')
```

Output:
```
<ipython-input-755-d1e69db46ca1> in <module>()
----> 1 check_report(res.stdout, 8, '$25.92')

<ipython-input-742-958bf15719e4> in check_report(text, expected_count, expected_total)
     10             for s in re.findall(r'\$\d+\.\d\d', line):
     11                 nums.append(s)
---> 12     assert(len(nums) == expected_count)
     13     total = 0.0
     14     for s in nums:

AssertionError:
```

##### <span style="color:red">Autograder Tests</span>

The tests for the `sakila_report` function will call the function and while using the IPython `capture` command to save everything the function prints.  We'll then use a function that counts the number of times a dollar amount (from either a rental fee or a late fee) occurs and make sure the total amount is the expected value.

**Do not edit, copy, or delete these cells.**

```python
import re

def check_report(text, expected_count, expected_total):
    nums = []
    for line in text.split('\n'):
        has_total = re.search(r'Total.*(\$\d+\.\d\d)', line)
        if has_total:
            assert has_total.group(1) == expected_total
        else:
            for s in re.findall(r'\$\d+\.\d\d', line):
                nums.append(s)
    assert(len(nums) == expected_count)
    total = 0.0
    for s in nums:
        total += float(s[1:])
    assert(round(total,2)) == round(float(expected_total[1:]), 2)
    return True
```

```python
%%capture res
sakila_report('Ebert', 6, 2005)
```

<!-- grade_id: leo_test, points: 4, grade: true, solution: false -->
```python
assert check_report(res.stdout, 8, '$25.92')
```

```python
%%capture res
sakila_report('Smith', 8, 2005)
```

<!-- grade_id: mary_test, points: 4, grade: true, solution: false -->
```python
assert check_report(res.stdout, 15, '$32.85')
```

```python
%%capture res
sakila_report('Turing', 6, 2005)
```

<!-- grade_id: alan_test, points: 2, grade: true, solution: false -->
```python
assert len(res.stdout) == 0
```

```python

```

```python

```